<a href="https://colab.research.google.com/github/Lakshmipriya-M-P/hds5210-2023/blob/main/week03/week03_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 3 Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/controlflow.html) section 4 on flow control.


**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 14.1 Difference in rate per 1,000

Often in public health, we report metrics as a number per 1,000 or per 1,000,000 in population. The purpose of that is to normalize the numbers between area of larger and smaller populations.

Below, we have a function already built to calculate teh different in rate per 1,000 in two different regions.  The function normalizes them and then prints out a message describing their difference.

```python
def diff_in_rate_normalized(count_a, total_a, count_b, total_b):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B

    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population."""

    rate_a = count_a / total_a * 1000
    rate_b = count_b / total_b * 1000

    if rate_a == rate_b:
        msg = "The rate in A and the rate in B are the same ({}).".format(rate_a)
    elif rate_a > rate_b:
        msg = "The rate in A ({}) is greater than the rate in B ({}).".format(rate_a, rate_b)
    else:
        msg = "The rate in A ({}) is less than the rate in B ({})".format(rate_a, rate_b)

    return msg
```

```python
rate_covid_testing = diff_in_rate_normalized(1,1000,15,10000)
print(rate_covid_testing)
```

Output:
```
The rate in A (1.0) is less than the rate in B (1.5)
```

```python
diff_in_rate_normalized(3,1000,30,10000)
```

Output:
```
'The rate in A and the rate in B are the same (3.0).'
```

### 14.2 Trimming outliers

While it works well in most cases, in practice, this may not truly make numbers comparable between extremely large population centers (e.g. New York City at 8.5 million) and very small rural areas (e.g. Meeteese, WY at 459 people).

Let's take that function and make some adjustments.  If the total population of either A or B is more than 100 times larger than the other, then we want to return a different message.  That is, if the populations are more than two orders of magnitude different, then we shouldn't try to compare them.

```python
def diff_in_rate_normalized(count_a, total_a, count_b, total_b):
    """(int,int,int,int) -> str
    * count_a is the number of occurences in area A
    * total_a is the total population in area A
    * count_b is the number of occurences in area B
    * total_b is the total population in area B

    This function returns a string describing how A and B compare in terms of occurence rates per 1,000 population.

    If total_a / total_b > 100 OR total_b / total_a > 100 then we'll return a message saying the two can't be compared.
    "The total populations in A and B are so different that they can't be compared."
    """

    msg=""

    if total_a/total_b > 100 or total_b/total_a>100:
        msg="The total populations in A and B are so different that they can't be compared."
        return msg

    rate_a = count_a / total_a * 1000
    rate_b = count_b / total_b * 1000

    if rate_a == rate_b:
        msg = "The rate in A and the rate in B are the same ({}).".format(rate_a)
    elif rate_a > rate_b:
        msg = "The rate in A ({}) is greater than the rate in B ({}).".format(rate_a, rate_b)
    else:
        msg = "The rate in A ({}) is less than the rate in B ({})".format(rate_a, rate_b)

    return msg
```

```python
diff_in_rate_normalized(5, 459, 30, 8500)
```

Output:
```
'The rate in A (10.893246187363834) is greater than the rate in B (3.5294117647058827).'
```

### 14.3 LACE Score
There is a simple readmission index called the LACE Score: https://www.hindawi.com/journals/bmri/2015/169870/tab1/

Use this documentation to create a function that can compute LACE scores based on the 4 input parameters.  Below is the function signature and documentation that you start with.

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def LACE(length_of_stay, acute_flag, charlson, ed_visits):
    if length_of_stay<=0:
        length_of_stay=0
    elif length_of_stay>=4 and length_of_stay<=6:
        length_of_stay=4
    elif length_of_stay>=7 and length_of_stay<=13:
        length_of_stay=5
    elif length_of_stay>=14:
        length_of_stay=7
    else:
        pass

    acute_flag = 3 if(acute_flag==True) else 0
    charlson= 5 if (charlson>=4) else charlson
    ed_visits= 4 if (ed_visits>=4) else ed_visits

    score= length_of_stay+acute_flag+charlson+ed_visits
    return score
```

<!-- grade_id: bmi-test, points: 10, grade: true, solution: false -->
```python
assert LACE(4, False, 1, 0) == 5
assert LACE(4, True, 4, 7) == 16
```

```python

```

### 14.4 Care Management Criteria

Care managers use LACE as part of the criteria for assigning a care coodinator to a patient who has been recently discharged. If the score is above 10, then a care coordinator will be assigned. The other criteria they use is if the patient has been discharged with a diagnosis of CHF or COPD.  If the diagnosis field has CHF or COPD in it, then the patient will have a care coordinator assigned.

For this exercise, write another function that takes the same inputs as LACE() plus another diagnosis parameter, and return True or False depending on if the paient needs a care coordinator.

**NOTE** Pay attention to the fact that the order of parameters in this function definition are not the same as the order they were in the LACE score.  

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def assign_care_coordinator(diagnosis_cd, ed_visits, length_of_stay, acute_flag, charlson):

    assignment = True
    if diagnosis_cd.upper()=="CHF" or diagnosis_cd.upper()=="COPD":
        return assignment
    else:
        pass

    if length_of_stay<=0:
        length_of_stay=0
    elif length_of_stay>=4 and length_of_stay<=6:
        length_of_stay=4
    elif length_of_stay>=7 and length_of_stay<=13:
        length_of_stay=5
    elif length_of_stay>=14:
        length_of_stay=7
    else:
        pass

    acute_flag = 3 if(acute_flag==True) else 0
    charlson= 5 if (charlson>=4) else charlson
    ed_visits= 4 if (ed_visits>=4) else ed_visits

    score= length_of_stay+acute_flag+charlson+ed_visits

    assignment= True if (score>10) else False
    return assignment


```

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
assert assign_care_coordinator('None', 0, 4, False, 1) == False
assert assign_care_coordinator('CHF', 0, 4, False, 1) == True
assert assign_care_coordinator('COPD', 0, 4, False, 1) == True
assert assign_care_coordinator('None', 7, 4, True, 4) == True
assert assign_care_coordinator('CHF', 7, 4, True, 4) == True
```

```python

```

### 14.5 qCSI COVID-19 Severity Index

See: https://www.mdcalc.com/quick-covid-19-severity-index-qcsi#evidence

Calculate the total risk score as per the point values assigned to respiratory rate, pulse oximetry, and O2 flow rate.  Then calculate and return the Risk Level.

In addition to the rules provided at the link above, also add the following checks for valid values:
* If `respiratory_rate <= 0` then return _invalid respiratory rate_
* If `pulse_ox <= 0` then return _invalid pulse ox_
* If `pulse_ox > 100` then return _invalid pulse ox_
* If `os_flow <= 0` then return _invalid O2 flow rate_

```python
def qcsi(respiratory_rate, pulse_ox, o2_flow):

    rtrn= "invalid respiratory rate" if (respiratory_rate <= 0) else "invalid pulse ox" if (pulse_ox <= 0 or pulse_ox > 100) else "invalid O2 flow rate" if (o2_flow <= 0) else 'valid'

    if rtrn!="valid":
        return rtrn
    else:
        pass

    respiratory_rate= 0 if (respiratory_rate <= 22) else 1 if ((respiratory_rate >= 23) and (respiratory_rate <= 28)) else 2 if (respiratory_rate > 28) else respiratory_rate

    pulse_ox= 0 if pulse_ox > 92 else 2 if ((pulse_ox >= 89) and (pulse_ox <= 92)) else 5 if pulse_ox <= 88 else pulse_ox

    o2_flow= 0 if o2_flow <=2 else 4 if ((o2_flow == 3) or (o2_flow == 4)) else 5 if ((o2_flow == 5) or (o2_flow == 6)) else o2_flow

    total=respiratory_rate+pulse_ox+o2_flow

    risk= "low" if total <=3 else 'low-intermediate' if ((total >= 4) and (total <= 6)) else 'high-intermediate' if ((total >= 7) and (total <= 9)) else 'high' if ((total >= 10) and (total <= 12)) else 'None'

    return risk
```

```python
assert (qcsi(29, 95, 1) == 'low')
assert (qcsi(20, 93, 1) == 'low')
assert (qcsi(29, 88, 1) == 'high-intermediate')
assert (qcsi(29, 88, 4) == 'high')
assert (qcsi(30, 90, 1) == 'low-intermediate')
assert (qcsi(28, 92, -1) == 'invalid O2 flow rate')
assert (qcsi(22, 0, 4) == 'invalid pulse ox')
assert (qcsi(0, 97, 2) == 'invalid respiratory rate')
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week03/week03_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**



---